# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before you start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
	# This example question requires counting words in the example_string below.
	example_string = "Amy is 5 years old"

	# YOUR CODE HERE.
	# You should write your solution here, and return your result, you can comment out or delete the
	# NotImplementedError below.
	result = example_string.split(" ")
	return len(result)

	#raise NotImplementedError()
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
# Dependencies.
import re

# Get names from a string.
def get_names():
	simple_string = """
		Amy is 5 years old, and her sister Mary is 2 years old. 
		Ruth and Peter, their parents, have 3 kids.
	"""

	return re.findall("[A-Z]\w*", simple_string)

get_names()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(get_names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
# Dependencies.
import re

# Get grades.
def get_grades():
	with open ("assets/grades.txt", "r") as file:
		grades = file.read()
		b_grades = re.findall("^(.*): B\s*$", grades, re.MULTILINE)
		return b_grades

print(get_grades())
print(len(get_grades()))
```

```python

```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(get_grades()) == 16
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {
	"host":"146.204.224.152", 
	"user_name":"feest6811", 
	"time":"21/Jun/2019:15:45:24 -0700",
	"request":"POST /incentivize HTTP/1.1"
}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
# Dependencies.
import re

# Convert logs into a list of dictionaries.
def convert_logs():
	# Create an empty list.
	logs = list()
	with open("assets/log-data.txt", "r") as file:
		for line in file:
			# Create an empty dictionary.
			log = dict()
			# Get the host, user_name, time, and request.
			host = re.findall("\d+\.\d+\.\d+\.\d+", line)
			user_name = re.findall(" - (.*?) ", line)
			time = re.findall(" \[(.*?)\] ", line)
			request = re.findall(" \"(.*?)\"", line)
			# Add values to the dictionary.
			log["host"] = host[0]
			log["user_name"] = user_name[0]
			log["time"] = time[0]
			log["request"] = request[0]
			# Append the dictionary to the list.
			logs.append(log)
	return logs

print(convert_logs())
print(len(convert_logs()))
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(convert_logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in convert_logs(), "Sorry, this item should be in the log results, check your formating"

```